## Assignment 1.2 Python Classes

https://docs.python.org/3/tutorial/classes.html

- The constructor is always written as a function called __init__()
- It must always take as its first argument a **reference to the instance** being constructed
- Single underscore before a member variable denotes it protected and double underscore denotes it private but it is just a convention

Create a Class Rectangle such that its objects are initialized with its length and breadth.
Make two class methods get_area and get_perimeter which calculates and return the area and circumference respectively

<!-- grade_id: cell-7b9cb67eee934ad9, points: 0, grade: false, solution: true -->
```python
class Rectangle():
# Your code here
    def __init__(self, length, breadth):
        self.length = length
        self.breadth = breadth

    def get_area(self):
    # Code here
        return self.length * self.breadth

    def get_perimeter(self):
    #Code here
        return 2 * (self.length + self. breadth)

# YOUR CODE HERE
    def __str__(self):
        return "Rectangle(length={}, breadth={})".format(self.length, self.breadth)
    
    # Operator overloading dunder methods
    def __eq__(self, other):
        return self.get_area() == other.get_area()
    
    def __gt__(self, other):
        return self.get_area() > other.get_area()

    def __lt__(self, other):
        return self.get_area() < other.get_area()

    def __ge__(self, other):
        return self.get_area() >= other.get_area()

    def __le__(self, other):
        return self.get_area() <= other.get_area()

    def __ne__(self, other):
        return self.get_area() != other.get_area()
```

<!-- grade_id: cell-513c46a34a22e61e, points: 1, grade: true, solution: false -->
```python
rec1 = Rectangle(4,10)
assert(rec1.get_area() == 40)
```

### Magic Methods or Dunder methods
[Python Special Methods](https://docs.python.org/3/reference/datamodel.html#special-method-names)

<!-- grade_id: cell-bcdd2ef93e91a97a, points: 0, grade: false, solution: true -->
```python
# Modify the Rectangle class to print object as "Rectangle(length=a, breadth=b)"
# where a and b are length and breadth 

# YOUR CODE HERE
print(rec1)
```

Output:
```
Rectangle(length=4, breadth=10)
```

<!-- grade_id: cell-962a216ebf3b2ffd, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-8d1bcf7bdbbedf55, points: 0, grade: false, solution: true -->
```python
# Modify rectangle class to compare rectangle objects (comparing with areas) 
# Operator overloading of rich comparisons = ,> ,< ,<= ,>= ,!= , 

# YOUR CODE HERE
rect1 = Rectangle(1, 3)
rect2 = Rectangle(1, 2)

print(rect1 == rect2)
print(rect1 > rect2)
print(rect1 < rect2)
print(rect1 >= rect2)
print(rect1 <= rect2)
print(rect1 != rect2)
```

Output:
```
False
True
False
True
False
True
```

<!-- grade_id: cell-5f2f866622093343, points: 1, grade: true, solution: false -->
```python
assert(Rectangle(3,10)==Rectangle(3,10))
assert(Rectangle(2,10)<Rectangle(3,10))
assert(Rectangle(3,10)!=Rectangle(2,10))
assert(Rectangle(3,10)>=Rectangle(3,10))
```

### Inheritance

Derive a class **Square** from **Rectangle** with one properties "length" (initialized through constructor) and able to calculate area, perimeter and comparison

<!-- grade_id: cell-92fe1c73a1a02948, points: 0, grade: false, solution: true -->
```python
class Square(Rectangle):

# YOUR CODE HERE
    def __init__(self, length):
        self.length = length
        self.breadth = length
```

<!-- grade_id: cell-28da33eb75674da3, points: 1, grade: true, solution: false -->
```python
assert(Square(9).get_area()==81)
assert(Square(10) != Square(5))
```