## Assignment 1.3 Numpy
Numpy is a library for the Python programming language, adding support for large, multi-dimensional arrays and matrices, along with a large collection of high-level mathematical functions to operate on these arrays. 
[Numpy Official Documentation](http://www.numpy.org/)

<!-- grade_id: cell-407e2d60461b23e6, points: 0, grade: false, solution: true -->
```python
# Import Numpy

# YOUR CODE HERE
import numpy as np
```

<!-- grade_id: cell-c5ae32b83b1d3fbc, points: 1, grade: true, solution: false -->
```python

```

NumPy’s main object is the homogeneous multidimensional array. It is a table of elements (usually numbers), all of the same type, indexed by a tuple of positive integers.

Go through [this quick tutorial](https://docs.scipy.org/doc/numpy/user/quickstart.html) to get started.

<!-- grade_id: cell-e21e5e2a23f7260c, points: 0, grade: false, solution: true -->
```python
# Find indices of non-zero elements from A as array of X-coordinates and Y-coordinates 
# and store it in variable "indices"
# nonzero function could be useful
A = np.array([[1,4,2],[0,5,7],[0,4,0],[1,2,0]])
indices = None
# YOUR CODE HERE
indices = np.nonzero(A)
print(indices)
```

Output:
```
(array([0, 0, 0, 1, 1, 2, 3, 3]), array([0, 1, 2, 1, 2, 1, 0, 1]))
```

<!-- grade_id: cell-4734df5d6d726bbe, points: 1, grade: true, solution: false -->
```python

```

### Random Number Generation
The use of randomness is an important part of the configuration and evaluation of machine learning algorithms.

From the random initialization of weights in an artificial neural network, to the splitting of data into random train and test sets, to the random shuffling of a training dataset in stochastic gradient descent, generating random numbers and harnessing randomness is a required skill.

Check out: [Numpy Random Routines](https://docs.scipy.org/doc/numpy-1.15.1/reference/routines.random.html)

Try to learn about
- Random Int, Float Generation
- Seed
- Normal or Gaussian Distribution
- Shuffle NumPy Array

<!-- grade_id: cell-d9de39c7298f8a33, points: 0, grade: false, solution: true -->
```python
# create a 20X10 matrix of random numbers using seed = 7 and store it in B
seed = 7
B = None
# YOUR CODE HERE
np.random.seed(7)
B = np.random.rand(20, 10)
print(B)
```

Output:
```
[[ 0.07630829  0.77991879  0.43840923  0.72346518  0.97798951  0.53849587
   0.50112046  0.07205113  0.26843898  0.4998825 ]
 [ 0.67923     0.80373904  0.38094113  0.06593635  0.2881456   0.90959353
   0.21338535  0.45212396  0.93120602  0.02489923]
 [ 0.60054892  0.9501295   0.23030288  0.54848992  0.90912837  0.13316945
   0.52341258  0.75040986  0.66901324  0.46775286]
 [ 0.20484909  0.49076589  0.37238469  0.47740115  0.36589039  0.83791799
   0.76864751  0.31399468  0.57262533  0.27604905]
 [ 0.45284293  0.35297837  0.65739946  0.37035108  0.45909298  0.71932412
   0.41299183  0.90642327  0.18045162  0.74111887]
 [ 0.42237404  0.42645357  0.63437987  0.5229062   0.41488598  0.00142688
   0.09226235  0.70939439  0.5243456   0.69616046]
 [ 0.95546832  0.68291385  0.05312869  0.30885268  0.59259469  0.23512041
   0.964971    0.94504822  0.84840088  0.472324  ]
 [ 0.84147671  0.13111064  0.30873366  0.46299639  0.7418472   0.48582523
   0.13687612  0.34353653  0.32442617  0.3004189 ]

```

<!-- grade_id: cell-fdf9824a430587c0, points: 1, grade: true, solution: false -->
```python

```

### Reshaping
Machine learning data is generally represented as numpy arrays. Many times such data needs to be reshaped into another shape. Such as a 2-Dimensional array might be required to be changed to 1-Dimensional.

<!-- grade_id: cell-765efc861e170c4d, points: 0, grade: false, solution: true -->
```python
# Reshape the matrix B into 10X20 matrix and store it in C
C = None
# YOUR CODE HERE
C = B.reshape(10, 20)
```

<!-- grade_id: cell-2b07046cbadab284, points: 1, grade: true, solution: false -->
```python
assert(C.shape == (10,20))
```

Perform Matrix multiplication between B and C
$$
Prod = B \cdot C
$$

<!-- grade_id: cell-fc5b00521b726684, points: 0, grade: false, solution: true -->
```python
Prod = None
# YOUR CODE HERE
Prod = np.dot(B, C)
```

<!-- grade_id: cell-3a3d1357d470fa3e, points: 1, grade: true, solution: false -->
```python

```

Calculate A squared
$$
A^2 = A \cdot A^T
$$

<!-- grade_id: cell-696b88e57b529140, points: 0, grade: false, solution: true -->
```python
A_sq = None
# YOUR CODE HERE
A_sq = np.dot(A, A.transpose())
```

<!-- grade_id: cell-f6ad93fe17f1e87b, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-367f5a191fdac5f1, points: 0, grade: false, solution: true -->
```python
#concatenation
A = np.array([[1,2],[3,4]]) 
B = np.array([[5,6],[7,8]]) 
# Concate A and B vertically
concatenated = None
# YOUR CODE HERE
concatenated = np.vstack((A, B))
```

<!-- grade_id: cell-2b56e9185f1bcea0, points: 1, grade: true, solution: false -->
```python

```

Write a function that given inputs X and y returns theta as follows:
$$
theta = (X^{T}X)^{-1}X^{T}y
$$
where,
- X is a ndarray of size M X N
- y is a ndarray of size M X 1
Please go through linalg routines of numpy to solve this problem.
https://docs.scipy.org/doc/numpy/reference/routines.linalg.html

This is Normal Equation for [Ordinary Least Square](https://en.wikipedia.org/wiki/Ordinary_least_squares).

<!-- grade_id: cell-0790ab535a2aad5c, points: 0, grade: false, solution: true -->
```python
def f(X, y):
    # YOUR CODE HERE
    theta = np.dot(np.dot(np.linalg.inv(np.dot(X.transpose(), X)), X.transpose()), y)
    return theta

X = np.array([[1, 2], [3, 4]])
y = np.array([[5], [6]])
f(X, y)
```

Output:
```
array([[-4. ],
       [ 4.5]])
```

<!-- grade_id: cell-8815f55c00de7ca5, points: 1, grade: true, solution: false -->
```python

```