# Assignment 1.4 Pandas
Pandas is an open source Python package providing fast, flexible, and expressive data structures and data manipulation designed to make working with “relational” or “labeled” data both easy and intuitive. It is the fundamental high-level building block for doing practical, real world data analysis in Python.
https://pandas.pydata.org/

<!-- grade_id: cell-462d89bbf25b800a, points: 0, grade: false, solution: true -->
```python
## Import numpy and pandas

# YOUR CODE HERE
import pandas as pd
import numpy as np
```

<!-- grade_id: cell-ab515cd97836235b, points: 1, grade: true, solution: false -->
```python

```

The two primary data structures of pandas are 
- Series (1-dimensional) and 
- DataFrame (2-dimensional)

Read this [10 min intro to get started](https://pandas.pydata.org/pandas-docs/stable/10min.html)

### Data Loading

You are given two csv files.  
- 'sales1.csv'
- 'sales2.csv'

Load both CSV as dataframes and Investigate the data and its types

<!-- grade_id: cell-15bfe6e4718647be, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
sales1 = pd.read_csv('sales1.csv')
sales2 = pd.read_csv('sales2.csv')
```

<!-- grade_id: cell-3c85ac1eab5f8b19, points: 0, grade: false, solution: true -->
```python
# Investigate the data 
# YOUR CODE HERE
sales1.describe()
```

Output:
```
Order ID   Units Sold  Unit Price   Unit Cost
count  1.000000e+03   994.000000  994.000000  994.000000
mean   5.496813e+08  5046.129779  262.162193  185.051308
std    2.571334e+08  2897.630250  216.150584  175.419323
min    1.029280e+08    13.000000    9.330000    6.920000
25%    3.280740e+08  2404.750000   81.730000   41.047500
50%    5.566097e+08  5184.000000  154.060000   97.440000
75%    7.696945e+08  7533.500000  421.890000  263.330000
max    9.955298e+08  9998.000000  668.270000  524.960000
```

```python
sales1.dtypes
```

Output:
```
Item Type      object
Order Date     object
Order ID        int64
Ship Date      object
Units Sold    float64
Unit Price    float64
Unit Cost     float64
dtype: object
```

```python
sales2.describe()
```

Output:
```
Order ID  Total Revenue    Total Cost
count  9.800000e+02   9.770000e+02  9.770000e+02
mean   5.503923e+08   1.334170e+06  9.412845e+05
std    2.575631e+08   1.486748e+06  1.163293e+06
min    1.029280e+08   2.043250e+03  1.416750e+03
25%    3.273675e+08   2.899409e+05  1.679699e+05
50%    5.599769e+08   7.763065e+05  4.678675e+05
75%    7.699941e+08   1.761479e+06  1.144786e+06
max    9.955298e+08   6.617210e+06  5.204978e+06
```

```python
sales2.dtypes
```

Output:
```
Order ID            int64
Region             object
Country            object
Order Priority     object
Total Revenue     float64
Total Cost        float64
dtype: object
```

### Merge Dataframes
Merge the two dataframes into a single dataframe "sales" based on order id such that new data consist of common order ID and all the other columns.
Also analyse the data in different columns.
Use [Merge function](https://pandas.pydata.org/pandas-docs/version/0.23.4/generated/pandas.merge.html)

<!-- grade_id: cell-82ee3339bf5f4570, points: 0, grade: false, solution: true -->
```python
sales = None
# YOUR CODE HERE
sales = pd.merge(sales1 ,sales2 ,left_on="Order ID",right_on="Order ID",how="inner")
sales
```

Output:
```
Item Type  Order Date   Order ID   Ship Date  Units Sold  \
0          Cosmetics  10/18/2014  686800706  10/31/2014      8446.0   
1         Vegetables   11/7/2011  185941302   12/8/2011         NaN   
2          Baby Food  10/31/2016  246222341   12/9/2016      1517.0   
3            *Cereal   4/10/2010  161442649   5/12/2010      3322.0   
4             Fruits   8/16/2011  645713555   8/31/2011      9845.0   
5            *Cereal  11/24/2014  683458888  12/28/2014      9528.0   
6            *Cereal    3/4/2015  679414975   4/17/2015      2844.0   
7            Clothes   5/17/2012  208630645   6/28/2012      7299.0   
8         Vegetables   1/29/2015  266467225    3/7/2015      2428.0   
9         Vegetables  12/24/2013  118598544   1/19/2014      4800.0   
10           Clothes  12/29/2015  451010930   1/19/2016      3012.0   
11            Snacks   2/27/2010  220003211   3/18/2010      2694.0   
12         Household  11/17/2016  702186715  12/22/2016      1508.0   
13   Office Suppl
```

<!-- grade_id: cell-52d67c318540c515, points: 1, grade: true, solution: false -->
```python
assert(sales[sales["Order ID"]==686800706]["Country"].values == "Libya")
assert(sales[sales["Order ID"]==686800706]["Item Type"].values == "Cosmetics")
```

### Data Cleaning
If you look at the data we can see some issues,
- Some of the Item Type start with " * " we want to remove the asterik in the column
- Some columns data are missing indicated by NaN (u can try isnull(), isna() function)
  - If Region or Country data is missing replace it with unknown
  - if other data are missing, remove it from "sales" dataframe and add it in a "missing" dataframe 

Reference for this section:
- [Working with missing data](https://pandas.pydata.org/pandas-docs/stable/missing_data.html)
- [Python Lambda Expressions](https://docs.python.org/3/tutorial/controlflow.html#lambda-expressions)
- [Pandas Apply Function](https://pandas.pydata.org/pandas-docs/version/0.18/generated/pandas.Series.apply.html)
- [Pandas FillNa](https://pandas.pydata.org/pandas-docs/version/0.17.0/generated/pandas.DataFrame.fillna.html)

<!-- grade_id: cell-36c46b4d53bf011c, points: 0, grade: false, solution: true -->
```python
# Investigate dataframe for Null (NA/NAN) and count the number of rows with NA/NAN
NA_count = 0
# YOUR CODE HERE
for row_count in sales.isnull().sum(axis = 1):
    if row_count > 0:
        NA_count += 1
        
NA_count
```

Output:
```
11
```

<!-- grade_id: cell-c3f7ee60f7aec5bf, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-04b3a523548b2603, points: 0, grade: false, solution: true -->
```python
# remove "*", NaN, add missing sales data to "missing" dataframe 
# Replace NaN in column "Region" with "unknown"

sales["Region"].fillna('unknown', inplace=True)
sales["Item Type"] = sales["Item Type"].apply(lambda x: x.replace('*', ''))
```

```python
missing = sales[(pd.isnull(sales["Units Sold"]))]
# Remove rows with no sales records
sales = sales.drop(sales[pd.isnull(sales["Units Sold"])].index)
```

<!-- grade_id: cell-7bc1b916b4472109, points: 1, grade: true, solution: false -->
```python
assert(sales[sales["Order ID"]==267066323]["Item Type"].values == "Cereal")
assert(sales[sales["Order ID"]==267066323]["Region"].values == "unknown")
```

Find the number of unique items available.

<!-- grade_id: cell-588bcfca26c28277, points: 0, grade: false, solution: true -->
```python
NumItems = 0
# YOUR CODE HERE
NumItems = sales.groupby('Item Type')['Item Type'].nunique().sum()
```

<!-- grade_id: cell-51b5cda4685cd3d0, points: 1, grade: true, solution: false -->
```python

```

### Grouping and Analysis 
Find out the statistical summary for ‘Unit Cost’ column for each Items.
Then extract the count, mean and std for "Cereal","Fruits","Vegetables","Baby Food" and "Meat" as a dataframe "Kitchen"
References:
- [Pandas Indexing](https://pandas.pydata.org/pandas-docs/stable/indexing.html)
- [Groupby](https://pandas.pydata.org/pandas-docs/stable/groupby.html)
- [Loc](https://pandas.pydata.org/pandas-docs/version/0.23.4/generated/pandas.DataFrame.loc.html)
- [filter](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.filter.html)

<!-- grade_id: cell-d5b95b853fd96b60, points: 0, grade: false, solution: true -->
```python
Kitchen = None
# YOUR CODE HERE
Kitchen = pd.DataFrame(sales.groupby('Item Type')['Item Type'].count())
Kitchen['count'] = Kitchen['Item Type']
Kitchen.columns
```

Output:
```
Index(['Item Type', 'count'], dtype='object')
```

<!-- grade_id: cell-986e659a273a364c, points: 1, grade: true, solution: false -->
```python
assert(Kitchen.loc["Cereal"]["count"]==77)
```

<!-- grade_id: cell-346f31cb4ea47c75, points: 0, grade: false, solution: true -->
```python
#Generate the new column ‘Total Profit’ for each orders (Total Profit = Total Revenue - Total Cost)
# YOUR CODE HERE
sales['Total Profit'] = sales['Total Revenue'] - sales['Total Cost']
```

<!-- grade_id: cell-6419595bc71f6a09, points: 1, grade: true, solution: false -->
```python
assert(sales[sales["Order ID"]==246222341]["Total Profit"].values == 145419.62)
```

<!-- grade_id: cell-aae9d532c4dabf17, points: 0, grade: false, solution: true -->
```python
# How many data sets are there with profit greater than 60%
ans = None
# YOUR CODE HERE
ans = len(sales[sales['Total Profit'] > 1.6 * sales['Total Cost']])
```

<!-- grade_id: cell-c053463e75781632, points: 1, grade: true, solution: false -->
```python

```